::: {.exercise}
Now we have all the tools to analyse the relation between GDP per capita and fertility in Asian countries.

Starting from the `gap` data, subset the data in 2010 for Asian countries.
Compute a new column, `gdp_per_capita` by dividing gdp by population.

To visualize the result, we've added a `ggplot` call to make a scatter plot --
have a look at the `ggplot` code, but we will dive into visualization in a later chapter so no worries
if it doesn't make much sense yet.
:::


```{webr}
#| setup: true
#| exercise: ex_mutate
library(dslabs)
library(tidyverse)
gap <- as_tibble(gapminder)
gap_asia <- filter(gap, continent == "Asia")
```
```{webr}
#| exercise: ex_mutate
library(tidyverse)
library(dslabs)
gap <- as_tibble(gapminder)

gap_asia_2010 <- filter(...)
gap_asia_2010 <- mutate(...)

ggplot(gap_asia_2010, aes(x=gdp_per_capita, y=life_expectancy)) +
  geom_point() +
  theme_classic() +
  labs(title="Relation between GDP and life expectancy in Asia in 2010",
       caption="Source: Gapminder", x="GDP per Capita", y="Life expectancy")

```

::: { .hint exercise="ex_mutate"}
::: { .callout-tip title="Hint:"}

- For filter, see the section on subsetting data above.
- For mutate, use the pattern `mutate(data, new_column = calculation)`
- To calculate gdp per capita,  simply use `/` to divide one column by the other.
- Remember that you don't need to quote column names.


:::
:::

::: {.solution exercise="ex_mutate"}
::: { .callout-important icon=false title="🎓 Solution:"}

- For filter, see the section on subsetting data above.
- For mutate, use `mutate(gap_asia_2010, gdp_per_capita = calculation)`
- For the calculation, use `gdp / population` to divide one column by the other.
- Remember that you don't need to quote column names.

```{webr}
#| exercise: ex_mutate
#| solution: true
library(tidyverse)
library(dslabs)
gap <- as_tibble(gapminder)
gap_asia_2010 <- filter(gap, year == 2010, continent == "Asia")
gap_asia_2010 <- mutate(gap_asia_2010, gdp_per_capita=gdp / population)
ggplot(gap_asia_2010, aes(x=gdp_per_capita, y=life_expectancy)) +
  geom_point() +
  theme_classic() +
  labs(title="Relation between GDP and life expectancy in Asia in 2010",
       caption="Source: Gapminder", x="GDP per Capita", y="Life expectancy")


```

:::
:::

```{webr}
#| exercise: ex_mutate
#| check: true

gradethis::grade_this({
  if (grepl("\\.\\.\\.", .user_code)) {
    gradethis::fail("It looks like you haven't filled in the blanks yet — replace the `...` with your code.")
  }

  .result <- .envir_result$gap_asia_2010

  if (is.null(.result) || !is.data.frame(.result)) {
    gradethis::fail("Make sure to store your result as `gap_asia_2010`.")
  }

  gradethis::fail_if(!"gdp_per_capita" %in% names(.result),
                     "Did you add a `gdp_per_capita` column using `mutate()`?")
  gradethis::fail_if(!all(.result$continent == "Asia"),
                     "Did you filter for `continent == \"Asia\"`?")
  gradethis::fail_if(!all(.result$year == 2010),
                     "Did you filter for `year == 2010`?")

  expected <- dplyr::filter(tibble::as_tibble(dslabs::gapminder),
                             year == 2010, continent == "Asia") |>
    dplyr::mutate(gdp_per_capita = gdp / population)

  .msg <- gradethis::code_feedback()
  if (is.null(.msg)) {
    gradethis::pass_if_equal(expected)
    gradethis::pass()
  } else {
    gradethis::pass_if_equal(expected,
      message = "Excellent work! Your solution works, but is different from ours. If you're curious, use `ⓘ Show solution` to compare your solution with ours (you may have to click `💡 Show hint` first).")
    gradethis::fail(.msg, hint = FALSE)
  }
})
```


::: {.challenge}

Although you're not expected to fully understand the `ggplot` code yet,
that's no reason to not play around with it:

- Try putting another variable on the y-axis, such as infant mortality or fertility
- Try changing the theme, e.g. to `theme_linedraw` or `theme_void`
- Try adding `fill=region` to the `aes` function

As a more serious challenge, let's consider the non-linear relation we see in the original plot:
it climbs very steeply first and then plateaus.
Can you use the `log10` function to compute the log of the gdp before plotting?
Does that make the relationship more linear?

:::
