::: {.exercise}
The code below shows the relation between infant mortality and fertility in 2010.
Change it to show the relation between GDP per capita (`gdp_percap`) on the x axis 
and life expectancy (`life_expectancy`) on the y axis instead, keeping the trend line.
:::

```{webr}
#| setup: true
#| exercise: ex_gdp_life
library(dslabs)
library(tidyverse)
library(ggthemes)
gap <- as_tibble(gapminder)
gap_2010 <- filter(gap, year == 2010)
```

```{webr}
#| exercise: ex_gdp_life
#| warning: false
gap_2010 |>
  mutate(gdp_percap = gdp / population) |>
  ggplot(mapping = aes(x = infant_mortality, y = fertility)) +
  geom_point() +
  geom_smooth(method = "lm")
```

::: { .hint exercise="ex_gdp_life"}
::: { .callout-tip title="Hint:"}
Change both the `x` and `y` arguments inside `aes()`: use `x = gdp_percap` and `y = life_expectancy`.
:::
:::

::: {.solution exercise="ex_gdp_life"}
::: { .callout-important icon=false title="🎓 Solution:"}
Change both the `x` and `y` arguments inside `aes()`: use `x = gdp_percap` and `y = life_expectancy`.

```{webr}
#| exercise: ex_gdp_life
#| solution: true
gap_2010 |>
  mutate(gdp_percap = gdp / population) |>
  ggplot(mapping = aes(x = gdp_percap, y = life_expectancy)) +
  geom_point() +
  geom_smooth(method = "lm")
```
:::
:::

```{webr}
#| exercise: ex_gdp_life
#| check: true

gradethis::grade_this({
  if (inherits(.result, "ggplot")) {
    gradethis::fail_if(
      rlang::as_label(.result$mapping$x) == "gdp",
      "The `gdp` column is the *total* GDP of a country, so big countries would dominate. Which column does the `mutate` create?")
    gradethis::fail_if(
      rlang::as_label(.result$mapping$y) == "fertility",
      "The x axis is right, but the y axis still shows fertility. Which column holds life expectancy?")
    gradethis::fail_if(
      rlang::as_label(.result$mapping$x) == "infant_mortality",
      "The y axis is right, but the x axis still shows infant mortality. Which column does the `mutate` create?")
  }

  {{< include exercises/_check.qmd >}}

})
```

::: {.challenge title="Interpreting the effect of GDP"}
Does the linear trend line fit the cloud of points well? 

How does that change your interpretation of the relation between GDP and life expectancy?

For what type of countries does extra income have the largest effect of life expectancy? 

What would a better way be to visualize this relation?
:::